```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# HW 5A: Sampling
## Due Date: Thursday, February 23rd, 11:59 PM

You must submit this assignment to Gradescope by the on-time deadline, Thursday, February 23rd, 11:59 PM. Please read the syllabus for the grace period policy. No late
submissions beyond the grace period will be accepted. While course staff is happy to help you if you encounter difficulties with submission, we may not be able to respond to last-minute requests for assistance (TAs need to sleep, after all!). **We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to reach out to staff for submission support. 

This is part of a two-part assignment. After completing this part ("Homework 5A"), please read all instructions carefully to combine manually graded questions from Homework 5A with Homework 5B, and submit your work to both the coding and written portals of Gradescope.

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your notebook.

**Collaborators**: *list collaborators here*

## This Assignment

The purpose of this assignment is for you to combine Python, math, and the ideas in Data 8 to draw some interesting conclusions. The methods and results will help build the foundation of Data 100.

## Score Breakdown
Question | Manual? | Points
--- | --- | --
1ai | No | 1
1aii | No | 1
1bi | No | 1
1bii | No | 1
1biii | No | 1
1biv | No | 1
2a | No | 2
2b | Yes | 2
2c | Yes | 3
2d | Yes | 2
Total | 3 | 15


### Initialize your environment

```python
import numpy as np
np.random.seed(200)
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
plt.style.use('fivethirtyeight')
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Data 100's Cutest Pet Dog Photo?

The Data 100 course staff wants to host a Cutest Pet Dog Photo contest for the Data C100/C200 class. Suppose the following three pet dog photos are in contention:

<img src='images/Venecia.jpg' width="180px" /> <img src='images/Mimi.jpg' width="170px" /> <img src='images/Beck.jpg' width="180px" />

**(From left to right) Venecia, Mimi, and Beck.**

Suppose that this contest will be held during the final exam, where all students will vote on the cutest pet dog photo. **Before the actual contest**, however, earlier in the semester the course staff would like to predict the results by surveying students in the class. This process is similar to polling that occurs before a political election.

In this section, you are going to explore different methods to sampling and data collection methods.

<br/><br/>


<hr style="border: 1px solid #fdb515;" />

## Question 1

<br><br>

---

### Question 1a

One of the Head TAs, Bella, would like to understand the class opinion before the contest actually takes place. In this coming week, Bella decides to survey all students before the contest by sending out an email via bCourses asking all students enrolled in Data C100/200 to chooce the cutest from the three pet choices. You may assume no other students/users receive the survey. Bella closes the survey 12 hours after sending it out.

**Part 1**. Which of the following is the population of interest? Assign your answer choice to `q1ai` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students only</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section<br/>
E. All students who fill out Bella's survey

```python
q1ai = ...
```

```python
grader.check("q1ai")
```


**Part 2**. Which of the following is the sampling frame? Assign your answer choice to `q1aii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students only</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section<br/>
E. All students who fill out Bella's survey

```python
q1aii = ...
```

```python
grader.check("q1aii")
```



**Part 3**. Which of the following is the sample? Assign your answer choice to `q1biii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students only</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section<br/>
E. All students who fill out Bella's survey

```python
q1biii = ...
```

```python
grader.check("q1aiii")
```

<br><br>

---

### Question 1b

In practice, we cannot get 100\% survey response rate, often because our population is too large, or because there is a time limit. In this case, very few students answer Bella's survey before she closes it.

To get more data to predict the answer to the original question ("Data 100's Cutest Pet Dog Photo?"), Bella decides on a different strategy: She conducts the pre-contest survey synchronously and in-person among her discussion section students that same week. Bella conducts the survey by presenting the following slide:

<img src='images/survey.jpg' width="600px" />

She then asks every student who attends discussion that week for their opinion on the cutest of the three pets. Assume that all of Bella's students show up for discussion that week and answer the survey.


**Part I:** Which of the following is the population of interest? Assign your answer choice to `q1bi` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section students<br/>
E. All students who fill out Bella's survey

```python
q1bi = ...
```

```python
grader.check("q1bi")
```


**Part II:** Which of the following is the sampling frame? Assign your answer choice to `q2bii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section students<br/>
E. All students who fill out Bella's survey

```python
q1bii = ...
```

```python
grader.check("q1bii")
```


**Part III:** Which of the following is the sample? Assign your answer choice to `q2bii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. Data C100 students</br>
C. Data C100/C200 students<br/>
D. Bella's discussion section students<br/>
E. All students who fill out Bella's survey

```python
q1biii = ...
```

```python
grader.check("q1biii")
```

**Part IV:** Which of the following best describes the sample? Assign your answer choice to `q1biv`  (e.g., `'A'`, `'B'`, etc.)..

A. Probability Sample<br/>
B. Simple Random Sample<br/>
C. Convenience Sample

```python
q1biv = ...
```

```python
grader.check("q1biv")
```

<br><br>

<hr style="border: 1px solid #fdb515;" />


## Question 2
Suppose that Bella was able to sample 50 students in her section, where the distribution of cutests pets is: **2% Venecia,  43% Mimi, and 57% Beck**. This is our **sample distribution.** Based on this result, Bella predicts that Beck will win the contest.

Let us fast-forward to the end of the semester, where the contest has actually taken place. Assume that after tallying every student's votes in the entire class, the true popularity of each pet dog photo is: **8% Venecia, 53% Mimi, and 39% Beck**. This is our **population distribution**. The true winner was actually Mimi!

---

### Question 2a

With the true population distribution we know that Mimi ultimately wins the Cutest Pet Dog Photo Contest, but from Bella's results, it was predicted that Beck would have won. Perhaps there was some underlying issue with Bella's discussion section sampling process...

Probability samples can help us quantify sampling bias and chance error. Put briefly, if we assume that a sample distribution was selected at random from a known population, then we can quantify how likely that sample is to have arisen due to random chance (**chance error**). If the difference in sample and population distributions is too great, then we suspect that the given sample has **bias** in how it was selected from the population.

Write one line of code that runs 1000 independent simulations, where each simulation finds the proportion of voters who voted **Mimi** for in a sample of size 50, selected uniformly at random from the true population. You may assume that the true population is large enough such that the sample is a random sample with replacement. The output `samples` should be an array with 1000 elements, each of which is the proportion of **Mimi** votes in that simulated sample.

**Hint**: 

* Use `np.random.multinomial` ([documentation](https://numpy.org/doc/stable/reference/random/generated/numpy.random.multinomial.html)).
* Use `array[:, i]` to select the ith column of a 2D numpy array
* Feel free to print out `samples` or `samples.shape` to check the output!

```python
samples = ...
samples
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2b

Create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice function of your choice to visualize `samples`, which is the simulated distribution of a sample Mimi vote. Include descriptive title and labels. An example is included below. The total area under the plot must be normalized to 1. Your plot may not match exactly to ours due to randomness of the data generating process.

<img src='images/q2b.png' width="500px" />

```python
...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2c

According to Bella's own 50-person sample, 43\% of her discussion section reported that they would vote for Mimi in the end-of-semester contest.

In the cell below, create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice to visualize Bella's sample statistic superimposed on the simulated sample distribution you plotted in the previous part. In other words, include a vertical line that passes through 43\%, a vertical line that passes through the mean fo the simulated sample distribution, and the simulated sample distribution itself. You should choose contrasting colors and include descriptive title, labels and a legend if needed. An example is included below.

<img src='images/q2c.png' width="500px" />

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br><br>

---

### Question 2d

Based on your analysis above, could Bella's result have arisen due to chance? If not, what could be a potential source of bias?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 5A!

### Submission Instructions

Below, you will see two cells.
* Running the first cell will automatically generate a PDF of your answers to all questions that need to be manually graded.
* Running the second cell will automatically generate a zip with your autograded answers.

**You are responsible for combining this resulting Homework 5A PDF with your answers to Homework 5B**, then submitting **both** the coding and written portions of Homework 5 to their respective Gradescope portals:
* **Homework 05 Coding**: Submit your Jupyter notebook zip file for Homework 5A,
which can be generated and downloaded from DataHub by using the `grader.export()`
cell provided below.
* **Homework 05 Written**: Submit a single PDF to Gradescope that contains both
(1) your answers to all manually graded questions from this Homework 5A Jupyter
Notebook, and (2) your answers to all questions in Homework 5B.

To receive credit on this assignment, **you must submit both your coding and written
portions to their respective Gradescope portals.** Your written submission (a single
PDF) can be generated as follows:

1. Access your answers to manually graded Homework 5A  questions in one of three ways:
      * <b>*Automatically create PDF (recommended)*</b>: Run the first cell below and download the generated PDF. This function will extract your response to the manually graded questions and put them on separate page. This process may fail if your answer is not properly formatted; if this is the case, check out common errors and solution described on Ed or follow either of the two ways described below.
      * *Manually download PDF*: If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by clicking on `File -> Save and Export Notebook As... -> PDF`. If you choose to go this route, you must take special care to ensure all appropriate pages are chosen for each question on Gradescope.
      * *Take screenshots*: If that doesn't work either, you can take screenshots of your answers (and your code if present) to manually graded questions and include them as images in a PDF. The manually graded questions are listed at the top of the Homework 1A notebook.
    
1. Answer the Homework 1B written questions.

1. Combine these two sets of answers together into the same PDF, and submit to the appropriate Gradescope written portal. You can use PDF merging tools, e.g., [Adobe Reader](https://get.adobe.com/reader/), [Smallpdf](https://smallpdf.com/merge-pdf) or [Apple Preview](https://support.apple.com/en-us/HT202945).

1. **Important**: When submitting on Gradescope, you **must tag pages to each question correctly** (it prompts you to do this after submitting your work). This significantly streamlines the grading process for our readers. Failure to do this may result in a score of 0 for untagged questions.

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
export_notebook("hw05.ipynb", filtering=True, pagebreaks=True)
if(path.exists('hw05.pdf')):
    display(HTML("Download your PDF <a href='hw05.pdf' download>here</a>."))
else:
    print("\n Pdf generation fails, please try the other methods described above")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1ai (points: 1)
```
>>> assert q1ai.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1aii (points: 1)
```
>>> assert q1aii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1aiii (points: 1)
```
>>> assert q1biii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1bi (points: 1)
```
>>> assert q1bi.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1bii (points: 1)
```
>>> assert q1bii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biii (points: 1)
```
>>> assert q1biii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biv (points: 1)
```
>>> assert q1biv.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q2a (points: 2)
```
>>> assert all(samples <= 1)
>>> assert all(samples >= 0)
>>> assert (len(samples) == 1000)
```
